# Statystyka matematyczna - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autora rozwiązania (imię, nazwisko i numer indeksu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Imię i nazwisko:
Numer indeksu:

---

# Kolowkium zaliczeniowe: 1 termin, grupa 1

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

## Dane do automatycznych testów

```python
# Dane do testów 1
test_data_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["X"])
test_data_1.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

```python
test_data_2 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=15, random_state=7), columns=["X"])
test_data_2.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

## Zadanie 1 [4 pkt]

Badania kliniczne pokazały, że po podaniu pewnego leku w 10% przypadków występują skutki uboczne wymagające hospitalizacji.

W pewnej klinice podano lek grupie 30 pacjentów. Jakie jest prawdopodobieństwo, że co najmniej 25 z tych osób nie będzie wymagało hospitalizacji?

Wynik zapisz do zmiennej `p`.

<!-- grade_id: cell-328b607aafafa4ce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
k=25
N=30
p=sp.stats.binomtest(k,N,p=0.1).pvalue
#raise NotImplementedError()
```

<!-- grade_id: cell-e646a3ab55f95f64, points: 4, grade: true, solution: false -->
```python
# Komórka testowa
assert isinstance(p, float)

```

## Zadanie 2 [6 pkt]

Przygotuj funkcję `variance_estimation()`, która będzie dokonywała estymacji przedziałowej wariancji dla danych wejściowych w postaci szeregu szczegółowego zgodnie z informacjami zawartymi w docstring funkcji oraz ze schematem z załączonego do notatnika zestawu wzorów (przypadek dla nieznanej wartości oczekiwanej populacji).

Uwaga! Zadanie dotyczy tradycyjnej estymacji, nie metody bootstrap.

<!-- grade_id: cell-00b25bd896944c4d, points: 0, grade: false, solution: true -->
```python
def variance_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i pojedynczą kolumnę "X" z wartościami próby losowej,
    alpha: float
        Wartość określająca poziom istotności.
    
    Returns
    -------
    var_low: float
        Dolna granica wyliczonego przedziału ufności.
    var_high: float
        Górna granica wyliczonego przedziału ufności.
    """
    L=len(data);
    #if population_mean == "unknown":
    S=np.var(test_data_1,ddof=1,axis=0);
    lower=((L-1)*S)/sp.stats.chi2.ppf(1-alpha/2,L-1);
    upper=((L-1)*S)/sp.stats.chi2.ppf(alpha/2,L-1);
    return lower.iloc[0],upper.iloc[0]
    # YOUR CODE HERE
#raise NotImplementedError()
variance_estimation(test_data_1, 0.1)
```

Output:
```
(0.018268125369359626, 0.05442967510189219)
```

<!-- grade_id: cell-60ea6c2bcb228b58, points: 6, grade: true, solution: false -->
```python
# Komórka testowa
result = variance_estimation(test_data_1, 0.1)
assert type(result)==tuple or type(result)==list
assert len(result)==2
assert isinstance(result[0], float)
assert isinstance(result[1], float)

```

## Zadanie 3 [6 pkt]

Przygotuj funkcję `mean_NHST()`, która będzie weryfikować hipotezę o wartości przeciętnej dla danych wejściowych w postaci szeregu szczegółowego zgodnie z informacjami zawartymi w docstring funkcji oraz ze schematem z załączonego do notatnika zestawu wzorów.

Uwaga! Zadanie dotyczy tradycyjnej weryfikacji, nie metody bootstrap.

<!-- grade_id: cell-146cf426bb90e71e, points: 0, grade: false, solution: true -->
```python
def mean_NHST(data, H0, H1, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i kolumnę "X" wynikami próby losowej.
    H0: float
        Wartość przeciętna przyjęta jako hipoteza zerowa.
    H1: str
        Postać hipotezy alternatywnej, przyjmuje wartości:
        - two-sided: wartość przeciętna jest różna od wartości przyjętej w H0,
        - less: wartość przeciętna jest mniejsza od wartości przyjętej w H0,
        - greater: wartość przeciętna jest większa od wartości przyjętej w H0.
    alpha: float
        Wartość określająca poziom istotności.
    
    Returns
    -------
    t: float
        Wyliczona na podstawie próby losowej wartość statystyki t.
    t_alpha: float
        Wartość statystyki t_alpha.
    H: int
        Wynik testu statystycznego, przyjmuje wartość:
        0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha,
        1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha.
    """
    L=len(data)
    srednia=np.mean(data)
    var=np.var(data,ddof=1,axis=0)
    t=((srednia-H0)/np.sqrt(var))*np.sqrt(L)
    t=t.iloc[0]
    H=0
    if H1=="two-sided":
        t_alpha=sp.stats.t.ppf(1-alpha/2,L-1)
        if abs(t)>=t_alpha:
            H=1
    if H1=="less":
        t_alpha=sp.stats.t.ppf(alpha,L-1)
        if t<=t_alpha:
            H=1
    if H1=="greater":
        t_alpha=sp.stats.t.ppf(1-alpha,L-1)
        if t>=t_alpha:
            H=1
    return t_alpha,t,H
    # YOUR CODE HERE
mean_NHST(test_data_2, 5.08, "two-sided", 0.05)
#raise NotImplementedError()
```

Output:
```
(2.1447866879169273, -1.934667783708126, 0)
```

<!-- grade_id: cell-aa9e98593265cfc8, points: 6, grade: true, solution: false -->
```python
# Komórka testowa
result = mean_NHST(test_data_2, 5.08, "two-sided", 0.05)
assert type(result)==tuple or type(result)==list
assert len(result)==3
assert isinstance(result[0], float)
assert isinstance(result[1], float)
assert isinstance(result[2], float) or isinstance(result[2], int)

```

## Zadanie 4 [12 pkt]

Zmienna `pulse` zawiera informacje na temat osób biorących udział w badaniu tętna przed i po wysiłku fizycznym (trzy "okrążenia" chodząc w górę i w dół po schodach):
 - Active - tętno (uderzenia na minutę) po ćwiczeniach
 - Rest - tętno spoczynkowe (uderzenia na minutę)
 - Sex - płeć
 - Hgt - wzrost (w calach)
 - Wgt - waga (w funtach)

```python
pulse = pd.read_csv("pulse_data.csv", index_col=0)
pulse
```

Output:
```
Active  Rest     Sex  Hgt  Wgt
0        97    78  Female   63  119
1        82    68    Male   70  225
2        88    62    Male   72  175
3       106    74    Male   72  170
4        78    63  Female   67  125
..      ...   ...     ...  ...  ...
227     105    85  Female   64  150
228      82    74  Female   66  124
229     102    81    Male   69  172
230      87    67    Male   68  170
231      81    62    Male   68  151

[232 rows x 5 columns]
```

### a)

W oparciu o dane ze zmiennej `pulse` przygotuj wykresy korelacyjne i kde (pairplot) zestawiające ze sobą poszczególne parametry liczbowe (osie X, Y) z podziałem na płeć osób biorących udział w badaniu (rozróżnione kolorem) i ustaw:
 - inną od domyślnej skalę kolorów (każda płeć ma przypisany inny kolor wykresu),
 - inne od domyślnych własności punktów (kształt, wielkość, przeźroczystość) dobrane tak, aby wykres był jak najbardziej czytelny.

<!-- grade_id: cell-c8281a216810df8a, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
```

### b)

W oparciu o dane ze zmiennej `pulse` przygotuj wykresy violinplot i stripplot wybranego przez siebie parametru liczbowego (oś Y) z podziałem na rozróżnioną kolorem płeć (oś X).

Wykresy zestaw na pojedynczej figurze (1x2 pola na wykres). Przygotowując figurę kieruj się wytycznymi z początku notatnika. Dodatkowo ustaw:
 - dla wykresu violinplot:
     - inny od domyślnego kolor wykresu,
 - dla wykresu stripplot:
     - inne od domyślnych rozmiar, kolor i przeźroczystość punktów (dobierz parametry tak, aby wykres był jak najbardziej przejrzysty).

<!-- grade_id: cell-da00623fe210a9ed, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[105], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

### c)

Załóżmy, że zmianna `pulse` zawiera pełną populację wyników ekperymentu. Wylosuj z niej i zapisz do zmiennej `sample` próbę losową zawierającą po 20 osób każdej płci wylosowanych bez zwracania.

<!-- grade_id: cell-fd24200c1c04a1d0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sampleM=pulse[pulse['Sex']=='Male'].sample(n=20,replace=False)
sampleF=pulse[pulse['Sex']=='Female'].sample(n=20,replace=False)
sample=pd.concat([sampleM,sampleF])
sample
```

<!-- grade_id: cell-aa13981f3d95fd2d, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(sample) == pd.DataFrame

```

### d)

Korzystając z dedykowanej do tego funkcji pakietu scipy dokonaj estymacji przedziałowej współczynnika korelacji liniowej między tętnem aktywnym a pasywnym uczestników eksperymentu na poziomie istotności 1-alpha = 0.9.

Wynik zapisz do zmiennej `r_ci` zawierającej listę lub krotkę z dolną i górną granicą przedziału ufności. Ustaw wartość argumentu funkcji odpowiadającego za ziarno generatora liczb losowych funkcji na 127.

<!-- grade_id: cell-b6f52487824e8671, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
r=np.corrcoef(pulse['Rest'],pulse['Active'])[0,1]
alpha=0.1
N=len(pulse['Rest'])
r_low=r-sp.stats.norm.ppf(1-alpha/2)*(1-r**2)/np.sqrt(N)
r_high=r+sp.stats.norm.ppf(1-alpha/2)*(1-r**2)/np.sqrt(N)
r_ci=(r_low,r_high)
r_ci
#raise NotImplementedError()
```

Output:
```
(0.535618021435404, 0.6727561547766148)
```

<!-- grade_id: cell-94f8e37e60f83338, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r_ci)==tuple or type(r_ci)==list
assert len(r_ci)==2
assert isinstance(r_ci[0], float)
assert isinstance(r_ci[1], float)

```

### e)

Korzystając z dedykowanej funkcji pakietu scipy zweryfikuj hipotezę, że tętna aktywne kobiet i mężczyzn biorących udział w eksperymencie mają taki sam rozkład statystyczny. W obliczeniach przyjmij poziom istotności 1 - alpha = 0.98.

Zapisz obliczoną wartość pvalue do zmiennej `pvalue`, oraz przygotuj zmienną `H` zawierającą wynik testu statystycznego (0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha, 1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha).

<!-- grade_id: cell-b644105f89af5513, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pvalue=sp.stats.ks_2samp(pulse[pulse['Sex']=='Male']['Active'],pulse[pulse['Sex']\
                         =='Female']['Active']).pvalue
alpha=0.02
pvalue
H=0
if pvalue<=alpha:
    H=1
H
#raise NotImplementedError()
```

<!-- grade_id: cell-2dd896b4ca1ac760, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert isinstance(pvalue, float)
assert isinstance(H, float) or isinstance(H, int)

```

```python

```